```python
# Initialize Otter
import otter
grader = otter.Notebook("filesExceptions.ipynb")
```

## Lecture Section

As statisticians, data scientists, researchers, academics, etc., we often turn to a programming language to work with data. This lecture will cover:
* How to load in text (.txt) and csv (.csv) files
* How to read/iterate through the files
* How to handle exceptions & errors

In the next lesson, we will learn how to manipulate the data we read/load in!

### Files

When we read in a file, we need to open it first. There are two common methods to do this, but both use the same `open()` function.

The first creates a `pointer` variable. This `pointer` *points* at the beginning of the file. We use `fp` as the variable name for the `file pointer`.

If we are only reading from our file, we open it with the argument `r`. If we are writing to the file, we open it with the argument `w`.

```python
fp = open("data/couple_sentences.txt", 'r')
```

This fp is **iterable**, just like a list or string, but it is **not** directly useful.

```python
print(fp) # not useful
```

Output:
```
<_io.TextIOWrapper name='data/couple_sentences.txt' mode='r' encoding='UTF-8'>
```

There are several ways to iterate through the pointer. We can do it line-by-line with `.read()`.

```python
fp.read()
```

Output:
```
'Hi! You found the message!'
```

Once we read the first line, the `pointer` moves to the second. Let's see if there is something there:

```python
fp.read()
```

Output:
```
''
```

Nope! How do we get it back to the beginning, then?

We use the **method** `.seek()` with an argument of `0`!

```python
fp.seek(0)
```

Output:
```
0
```

Now we can iterate through our file in a different way. We can use a `for`-loop!

```python
for line in fp:
    print(line)

fp.seek(0)
```

Output:
```
Hi! You found the message!
0
```

If you want to skip a line, you can use the `.readline()` method. It will return the skipped line, too!

```python
first_line = fp.readline()
first_line
```

Output:
```
'Hi! You found the message!'
```

Whenever we read or write files, we need to close them at the end. This is even more important when we write to one. You can find the file we have written to in the 'data' folder after the code has run.

```python
file_obj_out = open('data/write_ex.txt',"w")
file_obj_out.write("you're doing great!")
file_obj_out.close()
```

Another way of reading data is to do it with a `with` statement. `with` automatically closes the file for you, so it's often preferred. In the next example, we try to iterate through the file after the `with`, but it won't let us because it automatically closed!

The `.strip()` method is just removing extra white-space from the beginning and end of our line, if it exists.

```python
with open('data/write_ex.txt', 'r') as file:
    for line in file:
        print(line.strip())

for line in file:
    print(line.strip())

```

Output:
```
you're doing great!
ValueError                                Traceback (most recent call last)
Cell In[10], line 5
      2     for line in file:
      3         print(line.strip())
----> 5 for line in file:
      6     print(line.strip())

ValueError: I/O operation on closed file.
```

Everything above works for most file-types: not just `.txt`. There are libraries that make reading some file types, like `.csv` easier, though. Before we get there, we will talk about some specific `.txt` items.

#### `.txt`

We can loop through by word and character, too.


```python
fp.seek(0)
for line in fp:
    print(line)
    for character in line:
        print(character)

fp.seek(0)
```

Output:
```
Hi! You found the message!
H
i
!
 
Y
o
u
 
f
o
u
n
d
 
t
h
e
 
m
e
s
s
a
g
e
!
0
```

If we want to grab the word, we can use the `string` method `.split()` with a space ` ` as the argument!

```python
for line in fp:
    #print(line)
    for word in line.split(" "):
        print(word)
        for character in word:
            pass
            #print(character)

fp.seek(0)
```

Output:
```
Hi!
You
found
the
message!
0
```

If we want to grab the word, we can use the `string` method `.split()` with a space as the argument!

```python
for line in fp:
    #print(line)
    for word in line.split(" "):
        print(word)
        for character in word:
            pass
            #print(character)

fp.seek(0)
```

Output:
```
Hi!
You
found
the
message!
0
```

This method of reading data will also work for `.csv` files; however, it is not very efficient. You will need to split the data on the coma `,`, and it can get messy very quickly. If we try to do it without splitting, we don't get the results we expect.

```python
fp2 = open("data/random_data.csv", 'r')
fp2.read()
```

Output:
```
'x, 1, 3.5, "ye,llow"\ny, 2, 4.7, "bl,ue"\n'
```

#### `.csv`

With `.csv` files, we can use the `csv` library and the `csv.reader()` method. We still need to `open()` the file, though.

```python
import csv
fp = open("data/random_data.csv", 'r')
csv_reader = csv.reader(fp)
csv_reader
```

Output:
```
<_csv.reader at 0x722981b8b1b0>
```

We can iterate through the file like usual, now

```python
for row in csv_reader:
    print(row)
```

Output:
```
['x', ' 1', ' 3.5', ' "ye', 'llow"']
['y', ' 2', ' 4.7', ' "bl', 'ue"']
```

We have two problems here:

**All data is read through as a `string`, regardless of what it *should* be.**

This is true for all file-types; however, it's easier to forget when reading in `.csv` files. Always convert numerical data when you read it in.

Our other issue is with our colors. If you look in the original file, you'll see the colors are quoted but have a coma within them. They should not be split. Since quoted variables are accepted, and the quotes are removed, we just need to set the argument `skipinitialspace` in `csv.reader()` to `True`.

```python
fp.seek(0)
csv_reader = csv.reader(fp, skipinitialspace=True)
data_list = []
for row in csv_reader:
    data_list.append([row[0], int(row[1]), float(row[2]), row[3]])
print(data_list)

```

Output:
```
[['x', 1, 3.5, 'ye,llow'], ['y', 2, 4.7, 'bl,ue']]
```

If your file is deliminated by something other than a coma, which is the standard for coma-seperated csv files, you can change the deliminator argument in `csv.reader()`.

### Exceptions

This is a great time to introduce error-handling. Up to this point, we've only talked about avoiding exceptions/errors. Now, we will learn how to handle them!

When we try to open a file that does not exist, we get `FileNotFoundError`.

```python
fp = open("data/not_real.txt", 'r')
```

Output:
```
Cell In[18], line 1
----> 1 fp = open("data/not_real.txt", 'r')

File ~/.local/lib/python3.12/site-packages/IPython/core/interactiveshell.py:343, in _modified_open(file, *args, **kwargs)
    336 if file in {0, 1, 2}:
    337     raise ValueError(
    338         f"IPython won't let you open fd={file} by default "
    339         "as it is likely to crash IPython. If you know what you are doing, "
    340         "you can use builtins' open."
    341     )
--> 343 return io_open(file, *args, **kwargs)

FileNotFoundError: [Errno 2] No such file or directory: 'data/not_real.txt'
```

#### Try/Except

Luckily, we can use `try` and `except` to handle our error.

We begin with `try:`, and we put the code we are worried about beneath it.
After, we put `except:`, which is the code we run if an error or exception occurs in the code we `try:`

```python
try:
    fp = open("data/not_real.txt", 'r')
    fp.close()
except:
    fp = open("data/couple_sentences.txt", 'r')
    fp.close()
```

Your IDE might give you a warning about the `except:` above. This is because it will catch *any* error that occurs, and we should be more specific.

```python
try:
    fp = open("data/not_real.txt", 'r')
    fp.close()
except FileNotFoundError:
    fp = open("data/couple_sentences.txt", 'r')
    fp.close()
```

```python
try:
    y = x + 5
    fp = open("data/not_real.txt", 'r')
    fp.close()
except FileNotFoundError:
    fp = open("data/couple_sentences.txt", 'r')
    fp.close()

```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[21], line 2
      1 try:
----> 2     y = x + 5
      3     fp = open("data/not_real.txt", 'r')
      4     fp.close()

NameError: name 'x' is not defined
```

You still need to be careful - once an error is caught, the rest of the code under `try` won't run.

```python
try:
    fp = open("data/not_real.txt", 'r')
    fp.close()
    y = 7
except FileNotFoundError:
    fp = open("data/couple_sentences.txt", 'r')
    fp.close()
y
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[22], line 8
      6     fp = open("data/couple_sentences.txt", 'r')
      7     fp.close()
----> 8 y

NameError: name 'y' is not defined
```

If the code doesn't run, we won't catch other errors, either...

```python
try:
    fp = open("data/not_real.txt", 'r')
    fp.close()
    y = x + 5 # an error, but doesn't run!
except FileNotFoundError:
    fp = open("data/couple_sentences.txt", 'r')
    fp.close()
```

## Assignment Section

**Question 1.**

We will borrow a dataset from Kaggle for this exercise. I have edited it for the purposes of this course.
https://www.kaggle.com/datasets/msjahid/colorado-motor-vehicle-sales-data

For this problem, you will use `csv.reader()` to read the `"data/colorado_motor_vehicle_sales.csv"` file. You will extract the year [0], quarter [1], place [2] and sales [3] from each line and build a nested dictionary called `col_dict`. The dictionary format will be:

`col_dict[year (str) ] = {place (str) : sales (int)}`

If the place is equal to `xxx`, skip that line. If the sales cannot be converted to an `int` or are 0, skip that line. There will be more than one sale for most places and years: add the current sale to the one in the dictionary.

Hints:
1. Start by checking the skip conditions.
2. Check if year exists, then the place. If place exists, sales must exist.
3. If something doesn't exist, you know you need to add it!

Example:

`{ "2009" : {"Adams" : 387302, "Knox" : 823470}, "2010" : {"Adams" : 923740}}`





```python
import csv

fp = open("data/colorado_motor_vehicle_sales.csv", 'r')
csv_reader = csv.reader(fp)
col_dict = {}

for row in csv_reader:
    year = row[0].strip()
    quarter = row[1].strip()
    place = row[2].strip()
    sales_string = row[3].strip()

    if place == "xxx":
        continue
    
    try:
        sales = int(sales_string)
        if sales == 0:
            continue
    except ValueError:
        continue
    if year not in col_dict:
        col_dict[year] = {place: sales}       # add new year AND place
    elif place not in col_dict[year]:
        col_dict[year][place] = sales          # year exists, add new place
    else:
        col_dict[year][place] += sales         # both exist, add to sales
col_dict
```

Output:
```
{'2008': {'Adams': 840958000,
  'Arapahoe': 2092912000,
  'Boulder/Broomfield': 677201000,
  'Denver': 753352000,
  'Douglas': 350126000,
  'El Paso': 1202316000,
  'Fremont': 34561000,
  'Garfield': 206657000,
  'Jefferson': 1001485000,
  'La Plata': 90296000,
  'Larimer': 528820000,
  'Mesa': 351935000,
  'Pueblo': 259504000,
  'Weld': 575438000},
 '2009': {'Adams': 689914000,
  'Arapahoe': 1787961000,
  'Boulder': 471419000,
  'Broomfield': 106845000,
  'Denver': 714075000,
  'Douglas': 285998000,
  'El Paso': 1092017000,
  'Fremont': 30599000,
  'Garfield': 140116000,
  'Jefferson': 802371000,
  'La Plata': 67057000,
  'Larimer': 447041000,
  'Mesa': 253670000,
  'Pueblo': 234090000,
  'Weld': 425968000,
  'Rest of State': 103359000},
 '2010': {'Adams': 763549000,
  'Arapahoe': 1937961000,
  'Boulder': 515462000,
  'Broomfield': 120049000,
  'Denver': 730723000,
  'Douglas': 287024000,
  'El Paso': 1191001000,
  'Fremont': 34699000,
  'Garfield': 151032000,
  'Jefferson': 844852000
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2.** For this problem, you will create a function called `fix_list()`.
* Parameters: 1. A list.
* Returns: A new list.
* Goal: You are given a sample list to test the function. The function should iterate through a list, convert the items to an integer, then append the integer list items to new list, and return the new list. If a string has a single decimal in it, it needs to be converted to a float, first. If it has two decimals, remove the second decimal and convert. You are only converting ValueErrors. If the value is `None` or results in a TypeError, append 0 and move on.

```python
l = ["1", "1.2", "3..2", "4.3.2", "2001", None] # should be [1, 1, 3, 4, 2001, 0] when converted to integers
l2 = [None, '67..', '3', '..89']
def fix_list(l):
    new_l = []
    for i in range(len(l)):
        item = l[i]
        try:
            new_l.append(int(item))
        except TypeError:
            new_l.append(0)
        except ValueError:
            if item.count(".") == 1:
                new_l.append(int(float(item)))
            elif item.count(".") == 2:
                new_l.append(int(float(item[:item.rfind('.')] + item[item.rfind('.') + 1:])))
    return new_l

fix_list(l2)
```

Output:
```
[0, 67, 3, 0]
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3.** For this problem, you will create a function called `error_dictionary()`
* Parameters: 2. The first is a dictionary, and the second is a potential key value in the dictionary.
* Returns: Either "Found!" if the key variable is in the dictionary, or "Not Found!" if it is not.
* Example: `error_dictionary({'name':'Alex', 'color': 'green'}, 'music')` should return `Not Found!` since `music` is not a key in the given dictionary.

```python
def error_dictionary(dict_orig, key_value):
    try:
        dict_orig[key_value]
        return "Found!"
    except KeyError:
        return "Not Found!"

error = error_dictionary({'name': 'Alex', 'color': 'green'}, 'music')
print(error)

```

Output:
```
Not Found!
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!
```

Tests for q1 (points: 4)
```
>>> assert col_dict == {'2008': {'Adams': 840958000, 'Arapahoe': 2092912000, 'Boulder/Broomfield': 677201000, 'Denver': 753352000, 'Douglas': 350126000, 'El Paso': 1202316000, 'Fremont': 34561000, 'Garfield': 206657000, 'Jefferson': 1001485000, 'La Plata': 90296000, 'Larimer': 528820000, 'Mesa': 351935000, 'Pueblo': 259504000, 'Weld': 575438000}, '2009': {'Adams': 689914000, 'Arapahoe': 1787961000, 'Boulder': 471419000, 'Broomfield': 106845000, 'Denver': 714075000, 'Douglas': 285998000, 'El Paso': 1092017000, 'Fremont': 30599000, 'Garfield': 140116000, 'Jefferson': 802371000, 'La Plata': 67057000, 'Larimer': 447041000, 'Mesa': 253670000, 'Pueblo': 234090000, 'Weld': 425968000, 'Rest of State': 103359000}, '2010': {'Adams': 763549000, 'Arapahoe': 1937961000, 'Boulder': 515462000, 'Broomfield': 120049000, 'Denver': 730723000, 'Douglas': 287024000, 'El Paso': 1191001000, 'Fremont': 34699000, 'Garfield': 151032000, 'Jefferson': 844852000, 'La Plata': 72039000, 'Mesa': 256005000, 'Pueblo': 177402000, 'Rest of State': 436793000, 'Weld': 456412000, 'Larimer': 401355000}, '2011': {'Adams': 1376918000, 'Arapahoe': 2466086000, 'Boulder': 652528000, 'Broomfield': 144736000, 'Denver': 1089330000, 'Douglas': 422961000, 'El Paso': 1630360000, 'Fremont': 46555000, 'Garfield': 180631000, 'Jefferson': 1159715000, 'La Plata': 106218000, 'Larimer': 705870000, 'Mesa': 428593000, 'Pueblo': 231201000, 'Rest of State': 670836000, 'Weld': 775920000}, '2012': {'Adams': 853554000, 'Boulder': 656411000, 'Broomfield': 142109000, 'Denver': 817469000, 'Douglas': 362106000, 'El Paso': 1139875000, 'Fremont': 39345000, 'Garfield': 168111000, 'Jefferson': 1096652000, 'La Plata': 86476000, 'Larimer': 640163000, 'Mesa': 292028000, 'Pueblo': 284174000, 'Rest of State': 259680000, 'Weld': 622473000, 'Arapahoe': 1960170000}, '2013': {'Adams': 1174812000, 'Arapahoe': 2800854000, 'Boulder': 725073000, 'Broomfield': 163772000, 'Denver': 818590000, 'Douglas': 436626000, 'El Paso': 1237929000, 'Fremont': 37837000, 'Garfield': 218775000, 'Jefferson': 912764000, 'La Plata': 109278000, 'Larimer': 728581000, 'Mesa': 315673000, 'Pueblo': 312863000, 'Rest of State': 425655000, 'Weld': 687694000}, '2014': {'Adams': 1020084000, 'Arapahoe': 3113457000, 'Boulder': 828589000, 'Broomfield': 117331000, 'Denver': 896414000, 'Douglas': 494946000, 'El Paso': 1315028000, 'Fremont': 43544000, 'Garfield': 270844000, 'Jefferson': 1386771000, 'La Plata': 124291000, 'Larimer': 847686000, 'Mesa': 353035000, 'Pueblo': 343897000, 'Rest of State': 618013000, 'Weld': 554995000}, '2015': {'Adams': 1511503000, 'Arapahoe': 3408437000, 'Boulder': 893050000, 'Broomfield': 186445000, 'Denver': 943660000, 'Douglas': 596706000, 'El Paso': 1878968000, 'Fremont': 49799000, 'Garfield': 197166000, 'Jefferson': 1523003000, 'La Plata': 137152000, 'Larimer': 931283000, 'Mesa': 370497000, 'Pueblo': 392768000, 'Rest of State': 482269000, 'Weld': 776080000}}
```
```
>>> for year in col_dict.keys():
...     assert 'xxx' not in col_dict[year]
...     for place in col_dict[year]:
...         assert 0 != col_dict[year][place]
```


Tests for q2 (points: 4)
```
>>> l = ['1', '1.2', '3..2', '4.3.2', '2001', None]
>>> ans1 = [1, 1, 3, 4, 2001, 0]
>>> new_l = fix_list(l)
>>> assert len(new_l) == 6
>>> assert new_l == ans1
```
```
>>> l2 = [None, '67..', '3', '..89']
>>> ans2 = [0, 67, 3, 0]
>>> new_l2 = fix_list(l2)
>>> assert new_l2 == ans2
>>> assert len(new_l2) == 4
```


Tests for q3 (points: 2)
```
>>> error = error_dictionary({'name': 'Alex', 'color': 'green'}, 'music')
>>> assert error == 'Not Found!'
```
```
>>> error = error_dictionary({'name': 'Alex', 'color': 'green'}, 'name')
>>> assert error == 'Found!'
```
